## General instructions

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel/runtime** (Colab: in the menubar, select *Runtime*$\rightarrow$*Factory Reset Runtime*; Jupyter: in the menubar, select *Kernel*$\rightarrow$*Restart*) and then **run all cells** (Colab: in the menubar, select *Runtime*$\rightarrow$*Run all*; Jupyter: in the menubar, select *Cell*$\rightarrow$*Run All*).

Make sure you fill in any place that says `YOUR CODE HERE` or `"YOUR ANSWER HERE"`, as well as the list of the group members in the following cell.

Enter here the *Group Name* and the list of *Group Members*.

`Doi eletronics`

`Burco Lorenzo, Persello Riccardo`

In order to be able to have an evaluation DO NOT delete/cut the cells with code and answers. Once you have finished you can downolad the notebook (Colab: in the menubar, select *File*$\rightarrow$*Download .ipynb*; Jupyter: in the menubar, select *File*$\rightarrow$*Download as*$\rightarrow$*Notebook (.ipynb)*) and upload as an assignment on the e-learning platform.

The following cell will load the Google Drive extension for the current notebook, when the variable `MOUNT` is `True`. This allow you to mount the Google Drive filesystem for file persistence. The mountpoint will be `/content/gdrive`.
Furthermore, it will set the `PATH` variable, from now on, so that if you have to refer to external files you could do that by writing:

```python
os.path.join(PATH, filename)
```

This will append the filename after the specific PATH.

```python
import os
MOUNT = False
if 'google.colab' in str(get_ipython()) and MOUNT:
    from google.colab import drive
    drive.mount('/content/gdrive')
    PATH = '/content/gdrive/MyDrive'
else:
    PATH = '.'
```

# Important warning

**⚠️ avoid copying, removing or modifying test cells, if you do that your assignment might be graded wrongly ⚠️**

---

---

## Exercise 1

Using the repetition operator (`*`), create a list `zeros` that consists of 20 zeros, a list `ones` that consists of 20 replicates of the `set` containing only value `1` (i.e., `{1}`), and a list `twos` that consists of 20 replicates of the `tuple` consisting of value `2` (how would you write it?).

<!-- grade_id: lists_and_repetitions, points: 0, grade: false, solution: true -->
```python
"""
List creation
"""
zeros = [0] * 20
ones = [{1}] * 20
twos = [(2, )] * 20
print(zeros, ones, twos)
```

Output:
```
[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0] [{1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}, {1}] [(2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,)]
```

<!-- grade_id: correct_lists_and_repetitions, points: 1, grade: true, solution: false -->
```python
"""
  Check that the lists are correct.
"""
assert type(zeros) == list and type(ones) == list and type(twos) == list
for v in zeros:
    assert type(v) == int and v == 0
for v in ones:
    assert type(v) == set and v == {1}
for v in twos:
    assert type(v) == tuple and v == tuple([2])
```

## Exercise 2

Modify the fifth element of each list in the following way:

* for the `zeros` list set it to `-1`
* for the `ones` list add the value `3` to the set
* for the `twos` list concatenate the value `3` in the tuple

<!-- grade_id: fourth_change, points: 0, grade: false, solution: true -->
```python
"""
Modify the fifth element of each list
"""
zeros[4] = -1
ones[4].add(3)
twos[4] = (twos[5][0], 3)
print(zeros, ones, twos)
```

Output:
```
[0, 0, 0, 0, -1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0] [{1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}] [(2,), (2,), (2,), (2,), (2, 3), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,)]
```

<!-- grade_id: correct_fourth_change, points: 1, grade: true, solution: false -->
```python
"""Check that the modification is correct."""
assert zeros[4] == -1
assert 3 in ones[4]
assert 3 in twos[4]
# furthermore: enumerate will create an enumeration of index, value from the iterable
for i, v in enumerate(zeros):
    if i != 4:
        assert v == 0
for i, v in enumerate(ones):
    if i != 4:
        assert v == {1, 3}
for i, v in enumerate(twos):
    if i != 4:
        assert v == (2,)
```

The expected output after the modification is 
```python
[0, 0, 0, 0, -1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0] [{1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}, {1, 3}] [(2,), (2,), (2,), (2,), (2, 3), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,), (2,)]
```
How can you justify the different behavior in the three cases? Write down your answer in the following cell (try also using markdown syntax for enhancing the text).

<!-- grade_id: cell-8dd881f054bcafb7, points: 1, grade: true, solution: true -->
Sets behave differently because they are mutable containers. Integers and tuples are immutable. The array of sets simply contains twenty references to the same set, so modifying a single reference actually modifies the object referenced by all the elements.

Con you figure out an experiment with those containers to check you hypotheses for the previous question? Please write it down the code for verification.

<!-- grade_id: cell-2e41374b2a3d463d, points: 0, grade: true, solution: true -->
```python
x = {1, 2, 3}
y = x

# Modify y.
y.add(4)

# Check that x has been modified.
assert 4 in x

# Copy x.
z = x.copy()

# Modify z.
z.add(5)

# Check that x has not been modified.
assert 5 not in x
```

## Exercise 3

Given the `zeros` list, try to change all the values from index 4 to index 13 (included) with the integer values increasing from `1` to `10`. Do it using the *slicing* of the `zeros` list and the `range()` function for generating the numbers from `1` to `10`.

<!-- grade_id: slice_replacement, points: 0, grade: false, solution: true -->
```python
"""
Slice replacement
"""
zeros[4:13] = range(1, 11)
zeros
```

Output:
```
[0, 0, 0, 0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 0, 0, 0, 0, 0, 0, 0]
```

<!-- grade_id: correct_slice_replacement, points: 1, grade: true, solution: false -->
```python
"""
Check the slice replacement
"""
for i, v in enumerate(zeros):
    if i >= 4 and i <= 13:
        assert v == i - 4 + 1
```

---

# Word $n$-grams 

Given a string, a word $n$-gram is a subsequence of `n` consecutive words, all in lowercase and excluding punctuation. For example, the string `"He said bye as he walked through the door."` has the following (lowercased) word bi-grams (i.e., $n = 2$). 

```
[('he', 'said'), ('said', 'bye'), ('bye', 'as'), ('as', 'he'), ('he', 'walked'), ('walked', 'through'), ('through', 'the'), ('the', 'door')]
```

Similarly to the exercise in the previous practice we'd like to transform a sentence in a set of tuples of cardinality $n$ with the word $n$-grams.

## Exercise 4

Write a function `word_ngrams(s, n)` (possibly based on those function already defined in the previous practice), that given a string `s` and the size of the $n$-grams will output the set of $n$-grams of `s` represented as tuples. Notice, the `.split()` method of strings could be helpful.

<!-- grade_id: word_n_grams, points: 0, grade: false, solution: true -->
```python
import string

def remove_punctuation(s: str) -> str:
    """Removes the punctuation from the string s and returns the result"""
    for p in string.punctuation:
        s = s.replace(p, '')
    return s

def word_ngrams(s: str, n=2):
    """Returns the set of n-grams of the string"""

    words = remove_punctuation(s).lower().split(" ")

    result = set()
    for i in range(0, len(words)-(n-1)):
        result.add(tuple(words[i:i+n]))  

    return result

```

<!-- grade_id: correct_word_n_grams, points: 1, grade: true, solution: false -->
```python
assert word_ngrams("He said bye as he walked through the door.") == set([('he', 'said'),
 ('said', 'bye'),
 ('bye', 'as'),
 ('as', 'he'),
 ('he', 'walked'),
 ('walked', 'through'),
 ('through', 'the'),
 ('the', 'door')])
assert word_ngrams("He said bye as he walked through the door.", 3) == set([('he', 'said', 'bye'),
 ('said', 'bye', 'as'),
 ('bye', 'as', 'he'),
 ('as', 'he', 'walked'),
 ('he', 'walked', 'through'),
 ('walked', 'through', 'the'),
 ('through', 'the', 'door')])
```

```python

```

```python

```