## General instructions

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel/runtime** (Colab: in the menubar, select *Runtime*$\rightarrow$*Factory Reset Runtime*; Jupyter: in the menubar, select *Kernel*$\rightarrow$*Restart*) and then **run all cells** (Colab: in the menubar, select *Runtime*$\rightarrow$*Run all*; Jupyter: in the menubar, select *Cell*$\rightarrow$*Run All*).

Make sure you fill in any place that says `YOUR CODE HERE` or `"YOUR ANSWER HERE"`, as well as the list of the group members in the following cell.

Enter here the *Group Name* and the list of *Group Members*.

`Doi eletronics`

`Burco Lorenzo, Persello Riccardo`

In order to be able to have an evaluation DO NOT delete/cut the cells with code and answers. Once you have finished you can downolad the notebook (Colab: in the menubar, select *File*$\rightarrow$*Download .ipynb*; Jupyter: in the menubar, select *File*$\rightarrow$*Download as*$\rightarrow$*Notebook (.ipynb)*) and upload as an assignment on the e-learning platform.

The following cell will load the Google Drive extension, for the current notebook, allowing you to mount the Google Drive filesystem for file persistence. The mountpoint will be `/content/gdrive`, therefore you have to remember prefixing your file names with that path.

```python
if 'google.colab' in str(get_ipython()):
    from google.colab import drive
    drive.mount('/content/gdrive')
```

# Important warning

**⚠️ avoid copying, removing or modifying test cells, if you do that your assignment might be graded wrongly ⚠️**

---

The following is an example of manually graded answer, it is required to write down some code or some statement according to the type of question. Please tell me who are the group members and what's the name of the group in a *Markdown° cell.

<!-- grade_id: cell-978f414e0958f2d8, points: 1, grade: true, solution: true -->
```python

```

`Doi eletronics`

`Burco Lorenzo, Persello Riccardo`

A autograded answer, instead is a pair of cells that contain the code for the answer plus a test of your code that undergoes an automatic test. See the following example:

<!-- grade_id: cell-9b5ef44fe026141e, points: 0, grade: false, solution: true -->
```python
a = 42
```

<!-- grade_id: cell-9e14b46c56118d3f, points: 0, grade: true, solution: false -->
```python
assert a == 42, "Either the variable a does not exist or it has not been assigned the expected value"
```

```python

```