## General instructions

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel/runtime** (Colab: in the menubar, select *Runtime*$\rightarrow$*Factory Reset Runtime*; Jupyter: in the menubar, select *Kernel*$\rightarrow$*Restart*) and then **run all cells** (Colab: in the menubar, select *Runtime*$\rightarrow$*Run all*; Jupyter: in the menubar, select *Cell*$\rightarrow$*Run All*).

Make sure you fill in any place that says `YOUR CODE HERE` or `"YOUR ANSWER HERE"`, as well as the list of the group members in the following cell.

Enter here the *Group Name* and the list of *Group Members*.

`Doi eletronics`

`Burco Lorenzo, Persello Riccardo`

In order to be able to have an evaluation DO NOT delete/cut the cells with code and answers. Once you have finished you can downolad the notebook (Colab: in the menubar, select *File*$\rightarrow$*Download .ipynb*; Jupyter: in the menubar, select *File*$\rightarrow$*Download as*$\rightarrow$*Notebook (.ipynb)*) and upload as an assignment on the e-learning platform.

The following cell will load the Google Drive extension for the current notebook, when the variable `MOUNT` is `True`. This allow you to mount the Google Drive filesystem for file persistence. The mountpoint will be `/content/gdrive`.
Furthermore, it will set the `PATH` variable, from now on, so that if you have to refer to external files you could do that by writing:

```python
os.path.join(PATH, filename)
```

This will append the filename after the specific PATH.

```python
import os
MOUNT = False
if 'google.colab' in str(get_ipython()) and MOUNT:
    from google.colab import drive
    drive.mount('/content/gdrive')
    PATH = '/content/gdrive/MyDrive'
else:
    PATH = '.'
```

# Important warning

**⚠️ avoid copying, removing or modifying test cells, if you do that your assignment might be graded wrongly ⚠️**

---

---

## Exercise 1

Given the following list of tuples that represent a sequence of names and surnames, transform them into a *list* `person_list` whose elements are dictionaries with the following structure:

`{ 'name': <person name>, 'surname': <person surname>}`

```python
persons = [('Alexander', 'Thomas'),
 ('Monica', 'Grant'),
 ('美加子', '山本'),
 ('晃', '坂本'),
 ('Nicholas', 'White'),
 ('直人', '杉山'),
 ('直人', '山本'),
 ('Donald', 'Dickson'),
 ('Michelangelo', 'Giammusso'),
 ('Enzio', 'Battelli'),
 ('Jessica', 'Wade'),
 ('Fernanda', 'Paruta'),
 ('Dante', 'Comboni'),
 ('Arnulfo', 'Terragni'),
 ('陽子', '加納')]
```

<!-- grade_id: tuples_to_dict, points: 0, grade: false, solution: true -->
```python
"""
Convert the list of tuples into a list of dictionaries.
"""
person_list = []
for person in persons:
    person_dic = {'name': person[0], 'surname': person[1] }
    person_list.append(person_dic)

person_list
```

Output:
```
[{'name': 'Alexander', 'surname': 'Thomas'},
 {'name': 'Monica', 'surname': 'Grant'},
 {'name': '美加子', 'surname': '山本'},
 {'name': '晃', 'surname': '坂本'},
 {'name': 'Nicholas', 'surname': 'White'},
 {'name': '直人', 'surname': '杉山'},
 {'name': '直人', 'surname': '山本'},
 {'name': 'Donald', 'surname': 'Dickson'},
 {'name': 'Michelangelo', 'surname': 'Giammusso'},
 {'name': 'Enzio', 'surname': 'Battelli'},
 {'name': 'Jessica', 'surname': 'Wade'},
 {'name': 'Fernanda', 'surname': 'Paruta'},
 {'name': 'Dante', 'surname': 'Comboni'},
 {'name': 'Arnulfo', 'surname': 'Terragni'},
 {'name': '陽子', 'surname': '加納'}]
```

Your code should print `[{'name': 'Alexander', 'surname': 'Thomas'},
 {'name': 'Monica', 'surname': 'Grant'},
 {'name': '美加子', 'surname': '山本'},
 {'name': '晃', 'surname': '坂本'},
 {'name': 'Nicholas', 'surname': 'White'},
 {'name': '直人', 'surname': '杉山'},
 {'name': '直人', 'surname': '山本'},
 {'name': 'Donald', 'surname': 'Dickson'},
 {'name': 'Michelangelo', 'surname': 'Giammusso'},
 {'name': 'Enzio', 'surname': 'Battelli'},
 {'name': 'Jessica', 'surname': 'Wade'},
 {'name': 'Fernanda', 'surname': 'Paruta'},
 {'name': 'Dante', 'surname': 'Comboni'},
 {'name': 'Arnulfo', 'surname': 'Terragni'},
 {'name': '陽子', 'surname': '加納'}]` check that it does.

<!-- grade_id: correct_list, points: 1, grade: true, solution: false -->
```python
"""
  Check that squares returns the correct output for several inputs.
  The assert operator will check if a condition holds otherwise it fails
  Optionally a message can be provided after the condition so to inform about
  the error.
"""
assert person_list == [{'name': 'Alexander', 'surname': 'Thomas'},
 {'name': 'Monica', 'surname': 'Grant'},
 {'name': '美加子', 'surname': '山本'},
 {'name': '晃', 'surname': '坂本'},
 {'name': 'Nicholas', 'surname': 'White'},
 {'name': '直人', 'surname': '杉山'},
 {'name': '直人', 'surname': '山本'},
 {'name': 'Donald', 'surname': 'Dickson'},
 {'name': 'Michelangelo', 'surname': 'Giammusso'},
 {'name': 'Enzio', 'surname': 'Battelli'},
 {'name': 'Jessica', 'surname': 'Wade'},
 {'name': 'Fernanda', 'surname': 'Paruta'},
 {'name': 'Dante', 'surname': 'Comboni'},
 {'name': 'Arnulfo', 'surname': 'Terragni'},
 {'name': '陽子', 'surname': '加納'}], "The list is incorrect"
```

## Exercise 2

Write the Python code to count the number of characters (character frequency) in a string `s`. The code should build a dictionary, whose keys are the single characters in the string and the values are the number of occurences. The name of the variable that accumulates the result should be `count`.

<!-- grade_id: counting, points: 0, grade: false, solution: true -->
```python
"""
Count the occurences of each single character in a string `s`
"""
s = "google.com"
count = {}
for char in s:
    char_count = count.get(char, 0)
    count[char] = char_count+1

count
```

Output:
```
{'g': 2, 'o': 3, 'l': 1, 'e': 1, '.': 1, 'c': 1, 'm': 1}
```

The expected result for the string `"google.com"` is something like: `{'g': 2, 'o': 3, 'l': 1, 'e': 1, '.': 1, 'c': 1, 'm': 1}`.

<!-- grade_id: correct_count, points: 1, grade: true, solution: false -->
```python
"""Check that count is correct."""
assert count == {'g': 2, 'o': 3, 'l': 1, 'e': 1, '.': 1, 'c': 1, 'm': 1}
```

## Exercise 3

Write the Python code that removes all duplicates from a list `l`. For example if the list is `[1, 5, 1, 2, 4, 2]` the result might be `[1, 2, 4, 5]'`.

**Note**: there are several ways to solve this exercise.

<!-- grade_id: duplicates, points: 0, grade: false, solution: true -->
```python
"""
Remove the duplicates from list `l`
"""
l = [1, 5, 1, 1, 2, 1, 2, 4, 2, 4, 2]
s = set(l)
l = list(s)

l
```

Output:
```
[1, 2, 4, 5]
```

<!-- grade_id: correct_duplicates, points: 1, grade: true, solution: false -->
```python
"""
Check that there is no duplicate element (i.e., all elements' count is 1)
"""
assert type(l) == list, "l is not a list anymore"
for e in l:
    assert l.count(e) == 1


```

## Exercise 4

Given a string `s`, write the Python code that creates a new list `c` containing all the consonants of `s`. For example, if the string is `"google.com"` the expected result is `['g', 'g', 'l', 'c', 'm']`. The result must be stored in a list named `consonants`.

**Note**: to check if a character (or a whole string) is a letter you can use the method `.isalpha()`.

<!-- grade_id: consonants, points: 0, grade: false, solution: true -->
```python
"""
Extract the consonants from s
"""
s = "google.com"
consonants = []
for char in s:
    if char.isalpha() and char not in 'aeiou':
        consonants.append(char)

consonants
```

Output:
```
['g', 'g', 'l', 'c', 'm']
```

<!-- grade_id: cell-6d1c1c8d818ae9de, points: 0, grade: true, solution: false -->
```python
"""
Checks that the variab le consonants exists, it's a list, and contains the expected result
"""
try:
    consonants
except NameError:
    assert False, "The variable consonants does not exists"
assert type(consonants) == list
assert consonants == ['g', 'g', 'l', 'c', 'm']
```

---

# $n$-grams transformation

Given a string, an $n$-gram is a subsequence of `n` consecutive characters, all in lowercase, including spaces but excluding punctuation. For example, the string `"Ceramics collectables collectibles."` has the following bi-grams (i.e., $n = 2$). 

```
['ce', 'er', 'ra', 'am', 'mi', 'ic', 'cs', 's ', ' c', 'co', 'ol', 'll', 'le', 'ec', 'ct', 'ta', 'ab', 'bl', 'le', 'es', 's ', ' c', 'co', 'ol', 'll', 'le', 'ec', 'ct', 'ti', 'ib', 'bl', 'le', 'es']

```

What is relevant, from a data analysis point of view, is the frequency of the $n$-grams, that is how often they occur.

We aim at defining a set of functions that will lead to the transformation of the string into their $n$-grams frequency.

## Exercise 5

Write a function `remove_punctuation(s)` that given the string `s` returns a string that corresponds to `s` with removed punctuation. For help, refer to the documentation of the `string` library which can be imported through the `import string` statement (see [https://docs.python.org/3/library/string.html](https://docs.python.org/3/library/string.html) for details).

<!-- grade_id: punctuation, points: 0, grade: false, solution: true -->
```python
import string

def remove_punctuation(s: str):
    """Removes the punctuation from the string s and returns the result"""
    for p in string.punctuation:
        s = s.replace(p, '')
    return s
```

For example, the following statement should print out `googlecom`.

```python
remove_punctuation("google.com")
```

Output:
```
'googlecom'
```

<!-- grade_id: correct_punctuation, points: 1, grade: true, solution: false -->
```python
assert remove_punctuation("google.com") == "googlecom"
assert remove_punctuation(".,;-:") == ""
assert remove_punctuation("no punctuation") == "no punctuation"
assert remove_punctuation("no 3punct") == "no 3punct"
```

## Exercise 6

Write a function `divide(s, n)` that, given the string `s` and an integer `n`, whose default value should be $2$, splits the string into the list of $n$-grams of `s`. This means that the function should remove any punctuation and should normalize the $n$-grams so they will be in lowercase. The function *must* use the `remove_punctuation()` function defined before.

<!-- grade_id: divide, points: 0, grade: false, solution: true -->
```python
def divide(s: str, n=2):
    s = s.lower()
    s = remove_punctuation(s)
    result = []
    for i in range(0, len(s)-(n-1)):
        result.append(s[i:i+n])       
    return result
```

<!-- grade_id: correct_divide, points: 1, grade: true, solution: false -->
```python
assert divide("Ceramics collectables collectibles.", 2) == ['ce', 'er', 'ra', 'am', 'mi', 'ic', 'cs', 's ', ' c', 'co', 'ol', 'll', 'le', 'ec', 'ct', 'ta', 'ab', 'bl', 'le', 'es', 's ', ' c', 'co', 'ol', 'll', 'le', 'ec', 'ct', 'ti', 'ib', 'bl', 'le', 'es']
assert divide("Alpha Beta") == ['al', 'lp', 'ph', 'ha', 'a ', ' b', 'be', 'et', 'ta']
```

<!-- grade_id: correct_divide_2, points: 1, grade: true, solution: false -->
```python
"""
Checks that the function divide uses the function remove_punctuation
It removes the name from the environment (notice that is saved in orig_remove_punctuation)
and tries to invoke the `divide()` function expecting a failure. The `try-except-else` will be explained later.
"""
orig_remove_punctuation = remove_punctuation
del remove_punctuation
try:
    divide("Something", 2)
except NameError:
    pass
else:
    assert False, "The divide function do not use the remove_punctuation one"
finally:
    remove_punctuation = orig_remove_punctuation
```

## Exercise 7

Write a function `count_ngrams(s, n)` that, given the string `s` and the size `n` of the $n$-grams, returns a dictionary that counts the number of occurences of each $n$-gram of `s`. The function *must* use the previously defined one.

<!-- grade_id: ngrams, points: 0, grade: false, solution: true -->
```python
def count_ngrams(s, n=2):
    ngram_dict = {}
    for ngram in divide(s, n):
        count = ngram_dict.get(ngram, 0)
        ngram_dict[ngram] = count+1
    return ngram_dict
```

<!-- grade_id: correct_ngrams, points: 1, grade: true, solution: false -->
```python
assert count_ngrams("Ceramics collectables collectibles.", 2) == {'ce': 1, 'er': 1, 'ra': 1, 'am': 1, 'mi': 1, 'ic': 1, \
                                                                  'cs': 1, 's ': 2, ' c': 2, 'co': 2, 'ol': 2, 'll': 2, \
                                                                  'le': 4, 'ec': 2,  'ct': 2, 'ta': 1, 'ab': 1, 'bl': 2, \
                                                                  'es': 2, 'ti': 1, 'ib': 1}
assert count_ngrams("Ceramics collectables collectibles.") == {'ce': 1, 'er': 1, 'ra': 1, 'am': 1, 'mi': 1, 'ic': 1, \
                                                               'cs': 1, 's ': 2, ' c': 2, 'co': 2, 'ol': 2, 'll': 2, \
                                                               'le': 4, 'ec': 2,  'ct': 2, 'ta': 1, 'ab': 1, 'bl': 2, \
                                                               'es': 2, 'ti': 1, 'ib': 1}
assert count_ngrams("Ceramics collectables collectibles.", 3) == {'cer': 1, 'era': 1, 'ram': 1, 'ami': 1, 'mic': 1, \
                                                                  'ics': 1, 'cs ': 1, 's c': 2, ' co': 2, 'col': 2, \
                                                                  'oll': 2, 'lle': 2, 'lec': 2, 'ect': 2, 'cta': 1, \
                                                                  'tab': 1, 'abl': 1, 'ble': 2, 'les': 2, 'es ': 1, \
                                                                  'cti': 1, 'tib': 1, 'ibl': 1}
```

<!-- grade_id: correct_ngrams_use, points: 1, grade: true, solution: false -->
```python
"""
Checks that the function count_ngrams uses the function divide
"""
orig_divide = divide
del divide
try:
    count_ngrams("Something", 2)
except NameError:
    pass
else:
    assert False, "The count_ngrams function do not use the divide one"
finally:
    divide = orig_divide
```

## Exercise 8

Write a test that checks whether the dictionary keys returned by the `count_ngrams()` function have
the right size (e.g., if the size parameter `n` is set to 3 all of them are of length 3.

<!-- grade_id: ngrams_test_answer, points: 0, grade: false, solution: true -->
```python
s = "djaosidj asodijas saojdia"
for length in range(1, 10):
    for element in count_ngrams(s, length):
        assert len(element) == length, "An ngram doesn't respect the correct length"
```