## General instructions

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel/runtime** (Colab: in the menubar, select *Runtime*$\rightarrow$*Factory Reset Runtime*; Jupyter: in the menubar, select *Kernel*$\rightarrow$*Restart*) and then **run all cells** (Colab: in the menubar, select *Runtime*$\rightarrow$*Run all*; Jupyter: in the menubar, select *Cell*$\rightarrow$*Run All*).

Make sure you fill in any place that says `YOUR CODE HERE` or `"YOUR ANSWER HERE"`, as well as the list of the group members in the following cell.

Enter here the *Group Name* and the list of *Group Members*.

`Doi eletronics`

`Burco Lorenzo, Persello Riccardo`

In order to be able to have an evaluation DO NOT delete/cut the cells with code and answers. Once you have finished you can downolad the notebook (Colab: in the menubar, select *File*$\rightarrow$*Download .ipynb*; Jupyter: in the menubar, select *File*$\rightarrow$*Download as*$\rightarrow$*Notebook (.ipynb)*) and upload as an assignment on the e-learning platform.

The following cell will load the Google Drive extension for the current notebook, when the variable `MOUNT` is `True`. This allow you to mount the Google Drive filesystem for file persistence. The mountpoint will be `/content/gdrive`.
Furthermore, it will set the `PATH` variable, from now on, so that if you have to refer to external files you could do that by writing:

```python
os.path.join(PATH, filename)
```

This will append the filename after the specific PATH.

```python
import os
MOUNT = False
if 'google.colab' in str(get_ipython()) and MOUNT:
    from google.colab import drive
    drive.mount('/content/gdrive')
    PATH = '/content/gdrive/MyDrive'
else:
    PATH = '.'
```

# Important warning

**⚠️ avoid copying, removing or modifying test cells, if you do that your assignment might be graded wrongly ⚠️**

---

## Exercise 1

For memory efficiency, we do not want to store those values into a container but we aim at creating them in sequence when needed. For this purpose we want to define a generator function `prime_generator(n)` that will generate, one after another, the first $n$ prime numbers. That is, when the function is called for the next element it will `yield`s the next prime number.

```python
"""
Function prime
"""
def is_prime(n):
    if n == 0 or n == 1:
        return False
    for i in range(2, n // 2 + 1):
        if n % i == 0:
            return False
    return True 
```

<!-- grade_id: generator, points: 0, grade: false, solution: true -->
```python
"""
Fill the generator that yields the first n prime numbers.
The first two primes (2 and 3) are yielded directly so to possibly skip even numbers later in the code
"""
def prime_generator(n):
    if n <= 0:
        return None
    current = 2
    yield current
    if n > 1:
        current = 3
        yield current
        # the first two primes are 2 and 3, they are yielded directly so after we can skip even numbers
        for i in range(2, n):
            current += 1
            
            while not is_prime(current):
                current += 1
            
            yield current
list(prime_generator(10))
```

Output:
```
[2, 3, 5, 7, 11, 13, 17, 19, 23, 29]
```

<!-- grade_id: correct_generator, points: 1, grade: true, solution: false -->
```python
"""Check that the generator is correct."""
import types
assert type(prime_generator(1)) == types.GeneratorType, "You haven't defined a generator"
assert list(prime_generator(10)) == [2, 3, 5, 7, 11, 13, 17, 19, 23, 29]
for v in prime_generator(10):
    assert is_prime(v), f"{v} generated as a prime number but it is not"
```

<!-- grade_id: cell-8dd881f054bcafb7, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Exercise 2

Try to define a *fibonacci* generator, i.e. a generator that will generate a sequence of Fibonacci numbers (i.e., $0, 1, 1, 2, 3, 5, 8, 11, \ldots$). This will be a possibly *infinite* generator, meaning that it does not stop after a limit on the number of values generated and, therefore, it is not advisable to use it in a `list()` conversion expression. It is meant to be used, for example, for the search of a fibonacci number that is above a threshold value as in the following code:

```python
for f in fibonacci_generator():
    if f > threshold:
        break
```

<!-- grade_id: fibonacci_generator, points: 0, grade: false, solution: true -->
```python
"""
Fibonacci generator
Suggestions: 
- yield directly the first two fibonacci numbers and then keep the last number and the last-but-one 
number in variables to proceed on the sequence
- in order to create an infinite sequence you must use a `while True:` statement, notice that this does not mean
that enters an infinite loop, because the yield operator will stop the loop by yielding the values
"""
def fibonacci_generator():
    (lo, l) = (0, 1)
    yield lo
    yield l

    while True:
        yield l + lo
        (lo, l) = (l, l + lo)

    
for f in fibonacci_generator():
    if f > 400:
        break
print(f, "is the first fibonacci number greater than 400")
```

Output:
```
610 is the first fibonacci number greater than 400
```

<!-- grade_id: check_fibonacci_generator, points: 1, grade: true, solution: false -->
```python
"""Check that the generator is correct."""
import types
g = fibonacci_generator()
assert type(g) == types.GeneratorType, "You haven't defined a generator"
p, l = next(g), next(g)
assert p == 0 and l == 1, "One of those are not correct fibonacci numbers"
for i in range(1000): # check the generator for 1000 numbers
    f = next(g)
    assert f == p + l, f"{f} is not a correct fibonacci number"
    p, l = l, f
```

We want to make a textual analysis of the book **War and Peace** by *Leo Tolstoy*. To this aim you will find the book text in the file `book-war-and-peace.txt` which is distributed along with this notebook.

If you give a look at the file, you might notice that the format is the following one:

```
CHAPTER I

"Well, Prince, so Genoa and Lucca are now just family estates of the
Buonapartes. But I warn you, if you don't tell me that this means war,
if you still try to defend the infamies and horrors perpetrated by that
Antichrist--I really believe he is Antichrist--I will have nothing more
to do with you and you are no longer my friend, no longer my 'faithful
slave,' as you call yourself! But how do you do? I see I have frightened
you--sit down and tell me all the news."

It was in July, 1805, and the speaker was the well-known Anna Pavlovna
Scherer, maid of honor and favorite of the Empress Marya Fedorovna. With
these words she greeted Prince Vasili Kuragin, a man of high rank and
importance, who was the first to arrive at her reception. Anna Pavlovna
had had a cough for some days. She was, as she said, suffering from la
grippe; grippe being then a new word in St. Petersburg, used only by the
elite.
```

The paragraphs are split in multiple lines and there is one (or more) empty newline(s) among them but sentences can spread to multiple lines. Since we would like to specifically deal with sentences we want to recombine them so that they will represented as a whole.

# Exercise 3

Write a function `read_and_join(filename)` that, given a text file specified by its name with the brewvious format, will read the file and return it as a list of paragraphs (i.e., joining all the lines belonging to the same paragraph in a single list item). You can assume that the file always start with some non empty line, however you could not assume that there is just a single empty newline between paragraphs.

Suggestion: to detect whether a line is empty and consists only of a newline I suggest to use the string `.strip()` method and check whether the content is empty rather than comparing with `"\n"` (that's because there are different end of line customs for different operating systems).

Recall, also, to add a newline `\n` at the end of each paragraph.

<!-- grade_id: cell-14e8371f9f18871a, points: 0, grade: false, solution: true -->
```python
def read_and_join(filename):
    with open(filename, "r") as contents:
        lines = []
        current_line = ""

        for line in contents:
            if len(line.strip()) == 0 and len(current_line.strip()) != 0:
                current_line += "\n"
                lines.append(current_line)
                current_line = ""
            else:
                current_line += line.strip() + " "

        lines.append(current_line)

    return lines
```

<!-- grade_id: cell-28e6f4f08fe62e9c, points: 1, grade: true, solution: false -->
```python
book = read_and_join('./files/book-war-and-peace.txt')
assert len(book) == 132, 'the file contains 132 pars'
```

We want to measure the textual *similarity* of the different paragraphs of the book. We will speak more in detail about that in a future lecture, however up to now you will be guided through some steps to estimate this concept in different ways. 

First of all we need some way to represent those paragraphs in a compact and useful way. One possibility is to use the **Bag of words** *embedding* (BoW for short), in which a text (in general a whole *document*, however we will focus on paragraphs here) is represented as a *dictionary* whose keys are the words making up the text (normalized in lowercase and with punctuation removed) and the values are the frequencies of those words in the text.

For example, for the text `"John likes to watch movies. Mary likes movies too."` the BoW representation is `{"John": 1, "likes": 2, "to": 1, "watch": 1, "movies": 2, "Mary": 1, "too": 1}`. We have already seen something similar in the first practice when you were asked to count the $n$-grams instead. The idea is similar.

In order to create it, we will split into two phases.

## Exercise 4

Define a *generator* function `words(text)` that given a text will yield, one after the other, all the words that build up the text, removing all spaces and punctuation and converting them to lowercase.

For example, for the text `"John likes to watch movies. Mary likes movies too."` the generator should yield the sequence of words `john`, `likes`, `to`, `watch`, `movies`, `mary`, `likes`, `movies`, `too`. You can reuse (copy&paste) the functions that you're already define (such as remove_punctuation.

<!-- grade_id: cell-ccc277f80738ac49, points: 0, grade: false, solution: true -->
```python
import string

def remove_punctuation(s: str):
    """Removes the punctuation from the string s and returns the result"""
    for p in string.punctuation:
        s = s.replace(p, '')
    return s

def words(text):
    for word in remove_punctuation(text.lower()).strip().split(" "):
        yield word
    
list(words("John likes to watch movies. Mary likes movies too."))
```

Output:
```
['john', 'likes', 'to', 'watch', 'movies', 'mary', 'likes', 'movies', 'too']
```

<!-- grade_id: cell-9bf151f2cc84cc90, points: 0, grade: true, solution: false -->
```python
import string

for word in words("John likes to watch movies. Mary likes movies too."):
    for punct in string.punctuation:
        assert punct not in word, "there is some punctuation remained"
    assert word.strip() == word, "there is some whitespace in the word"
    assert word.lower() == word, "the word is not corresponding to the lowercase transformation"
assert list(words("John likes to watch movies. Mary likes movies too.")) == ['john', 'likes', 'to', 'watch', 'movies', 'mary', 'likes', 'movies', 'too']
assert list(words(book[42])) == ['attendez', 'said', 'anna', 'pavlovna', 'reflecting', 'ill', 'speak', 'to', 'lise', 'young', 'bolkonskis', 'wife', 'this', 'very', 'evening', 'and', 'perhaps', 'the', 'thing', 'can', 'be', 'arranged', 'it', 'shall', 'be', 'on', 'your', 'familys', 'behalf', 'that', 'ill', 'start', 'my', 'apprenticeship', 'as', 'old', 'maid']
```

Now that you defined the generator, we can use a `Counter`, which is a standard class available in the `collections` library. You can refer to [https://docs.python.org/3.8/library/collections.html#collections.Counter](https://docs.python.org/3.8/library/collections.html#collections.Counter) for some examples, however essentially it gets a list of immutable elements (as strings are), which can be a single string, and automatically gathers all the frequency of the single items, giving us automatically the BoW.

**Note:** you do not have to do nothing with this, just look at the example of combining the solution. If you want to engineer it you can wrap the creation of the BoW representation into a `bow(text)` function, but it's up to you.

```python
from collections import Counter

counts = Counter(words("John likes to watch movies. Mary likes movies too."))
print(counts)
```

Output:
```
Counter({'likes': 2, 'movies': 2, 'john': 1, 'to': 1, 'watch': 1, 'mary': 1, 'too': 1})
```

# Exercise 5

Try to create the BoW of the second and the fourth paragraph of the *Word and Peace* book.

<!-- grade_id: cell-ed491d321169dfd1, points: 0, grade: false, solution: true -->
```python
second = Counter(words(book[1]))
fourth = Counter(words(book[3]))
```

<!-- grade_id: cell-620a0a2bd5e64fc6, points: 1, grade: true, solution: false -->
```python
assert dict(second) == {'well': 1, 'prince': 1, 'so': 1, 'genoa': 1, 'and': 4, 'lucca': 1, 'are': 2, 'now': 1, 'just': 1, 'family': 1, 'estates': 1, 'of': 1, 'the': 3, 'buonapartes': 1, 'but': 2, 'i': 3, 'warn': 1, 'you': 7, 'if': 2, 'dont': 1, 'tell': 2, 'me': 2, 'that': 2, 'this': 1, 'means': 1, 'war': 1, 'still': 1, 'try': 1, 'to': 2, 'defend': 1, 'infamies': 1, 'horrors': 1, 'perpetrated': 1, 'by': 1, 'antichristi': 2, 'really': 1, 'believe': 1, 'he': 1, 'is': 1, 'will': 1, 'have': 2, 'nothing': 1, 'more': 1, 'do': 3, 'with': 1, 'no': 2, 'longer': 2, 'my': 2, 'friend': 1, 'faithful': 1, 'slave': 1, 'as': 1, 'call': 1, 'yourself': 1, 'how': 1, 'see': 1, 'frightened': 1, 'yousit': 1, 'down': 1, 'all': 1, 'news': 1}
assert dict(fourth) == {'all': 1, 'her': 1, 'invitations': 1, 'without': 1, 'exception': 1, 'written': 1, 'in': 1, 'french': 1, 'and': 1, 'delivered': 1, 'by': 1, 'a': 1, 'scarletliveried': 1, 'footman': 1, 'that': 1, 'morning': 1, 'ran': 1, 'as': 1, 'follows': 1}
```

## Exercise 6

One first possible similarity between two BoW is the degree of similarity between the two BoW seen as *sets* (i.e., disregarding the counts). This concept is called *Jaccard* similarity between sets $A$, and $B$ and is defined as follows:

$$J(A,B) = {{|A \cap B|}\over{|A \cup B|}}$$

Practically, the set representation of a text is just the set of *keys* of the BoW Counter (which is in fact a dictionary).

Define a function `jaccard(text_a, text_b)`, that given two texts transforms them into their set representation (passing through the BoW intermediate representation) and computes the Jaccard similarity between them.

Recall that the size of any container (including sets) can be accessed through the `len()` function.

<!-- grade_id: cell-0c768dfc18a844aa, points: 0, grade: false, solution: true -->
```python
def jaccard(text_a, text_b):
    a = set(words(text_a))
    b = set(words(text_b))

    return len(a & b)/len(a | b)

jaccard(book[42], book[3])
```

Output:
```
0.058823529411764705
```

<!-- grade_id: cell-a35f670422dbeb83, points: 0, grade: true, solution: false -->
```python
assert jaccard("AI is our friend and it has been friendly", "AI and humans have always been friendly") == 1/3
```

## Exercise 7

Another possibility for measuring the similarity is to vectorize the BoW of the two texts, so that the words are represented in a vector whose indexes corresponds to the words themseleves and the values to the counts of those word. 

However, for doing so, the indexes should univoquely refer to a specific word. To this aim the two sets of words in the two texts should be arbitrarily ordered. 

For example, when the two texts are 

```
John likes to watch movies. Mary likes movies too.
```
and
```
Mary also likes to watch football games.
```

The two BoWs are:

```
Counter({'likes': 2, 'movies': 2, 'john': 1, 'to': 1, 'watch': 1, 'mary': 1, 'too': 1})
```
and
```
Counter({'mary': 1, 'also': 1, 'likes': 1, 'to': 1, 'watch': 1, 'football': 1, 'games': 1})
```

For vectorizing the BoWs you must take the union of the two sets of words:

```
{'games', 'football', 'likes', 'to', 'mary', 'movies', 'also', 'too', 'john', 'watch'}
```

and transform them into a list (i.e., order it, although it can be either an arbitrary order or the words can be sorted), for example:

```
['also', 'football', 'games', 'john', 'likes', 'mary', 'movies', 'to', 'too', 'watch']
```

This means that `'also'` corresponds to index 0, `'football'` to index 1, etc.

This way, the two BoWs can be transformed in the following vectors (ordered):

```
[0, 0, 0, 1, 2, 1, 2, 1, 1, 1], 
[1, 1, 1, 0, 1, 1, 0, 1, 0, 1]
```

Notice, for example, that there are two zeros in the first vector stating that there is no occurence of words `'also'` and `'football'` in the first text, while the fifth value is `2` that corresponds to the occurences of `'likes'`.

Try define a function `vectorize_bow(text_a, text_b)` that, given the two texts performs all the needed operations to transform those texts into the vectorized Bag of Word which must be both returned by the function.

<!-- grade_id: cell-c71ef8b5b2334c0e, points: 0, grade: false, solution: true -->
```python
def vectorize_bow(text_a, text_b):
    wa = list(words(text_a))
    wb = list(words(text_b))

    all_words = wa + wb
    all_words.sort()

    empty_bow = { word : 0 for word in all_words }

    bow_a = empty_bow.copy()
    bow_a.update(Counter(wa))

    bow_b = empty_bow
    bow_b.update(Counter(wb))

    return (list(bow_a.values()), list(bow_b.values()))
```

Output:
```
([0, 0, 0, 1, 2, 1, 2, 1, 1, 1], [1, 1, 1, 0, 1, 1, 0, 1, 0, 1])
```

<!-- grade_id: cell-cd88ee8477a05915, points: 0, grade: true, solution: false -->
```python
assert vectorize_bow("John likes to watch movies. Mary likes movies too.", "Mary also likes to watch football games.") == ([0, 0, 0, 1, 2, 1, 2, 1, 1, 1], [1, 1, 1, 0, 1, 1, 0, 1, 0, 1])
```

## Exercise 8

Once you have the vectorized BoWs, you can use the so-called *cosine* similarity between the two vectors, which can be computed as follows:

$${\mathbf{a} \cdot \mathbf{b} \over \|\mathbf{a}\| \|\mathbf{b}\|}$$

where the $\cdot$ is the scalar product between the two vectors.

Define a function that given two texts computes their cosine similarity.

<!-- grade_id: cell-9ee4f5287c470ee6, points: 0, grade: false, solution: true -->
```python
from numpy import dot
from numpy.linalg import norm

sentence1 = "John likes to watch movies. Mary likes movies too."
sentence2 = "Mary also likes to watch football games."

def cosine(a, b):
    return dot(a, b) / (norm(a) * norm(b))

vect = vectorize_bow(sentence1, sentence2)
print(cosine(vect[0],vect[1]))
```

Output:
```
0.5241424183609592
```

<!-- grade_id: cell-bde21b0707b2c069, points: 1, grade: true, solution: false -->
```python
vect = vectorize_bow(sentence1, sentence2)
assert cosine(vect[0],vect[1]) == 0.5241424183609592
```

```python

```