## General instructions

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel/runtime** (Colab: in the menubar, select *Runtime*$\rightarrow$*Factory Reset Runtime*; Jupyter: in the menubar, select *Kernel*$\rightarrow$*Restart*) and then **run all cells** (Colab: in the menubar, select *Runtime*$\rightarrow$*Run all*; Jupyter: in the menubar, select *Cell*$\rightarrow$*Run All*).

Make sure you fill in any place that says `YOUR CODE HERE` or `"YOUR ANSWER HERE"`, as well as the list of the group members in the following cell.

Enter here the *Group Name* and the list of *Group Members*.

`Doi eletronics`

`Burco Lorenzo, Persello Riccardo`

In order to be able to have an evaluation DO NOT delete/cut the cells with code and answers. Once you have finished you can downolad the notebook (Colab: in the menubar, select *File*$\rightarrow$*Download .ipynb*; Jupyter: in the menubar, select *File*$\rightarrow$*Download as*$\rightarrow$*Notebook (.ipynb)*) and upload as an assignment on the e-learning platform.

The following cell will load the Google Drive extension for the current notebook, when the variable `MOUNT` is `True`. This allow you to mount the Google Drive filesystem for file persistence. The mountpoint will be `/content/gdrive`.
Furthermore, it will set the `PATH` variable, from now on, so that if you have to refer to external files you could do that by writing:

```python
os.path.join(PATH, filename)
```

This will append the filename after the specific PATH.

```python
import os
MOUNT = False
if 'google.colab' in str(get_ipython()) and MOUNT:
    from google.colab import drive
    PATH = '/content/gdrive'
    drive.mount(PATH)
else:
    PATH = '.'
```

# Important warning

**⚠️ avoid copying, removing or modifying test cells, if you do that your assignment might be graded wrongly ⚠️**

---

---

## Exercise 1

 Write one line of Python (i.e., use comprehensions) that takes a list a and makes a new list that has only the even elements of this list in it.

<!-- grade_id: cell-97108d46c8a8ee9d, points: 0, grade: false, solution: true -->
```python
def only_even(input_list):
    return [ x for x in input_list if x % 2 == 0 ]
    
only_even([1])
```

Output:
```
[]
```

<!-- grade_id: cell-c52f0c982f03893d, points: 1, grade: true, solution: false -->
```python
assert only_even([1,2,3,4,5,6,7,8,9]) == [2,4,6,8]
assert only_even([4,2,6,7]) == [4,2,6]
assert only_even([]) == []
assert only_even([1]) == []
```

---

## Exercise 2

 Write one line of Python (i.e., use comprehensions) that takes a list a returns the same list with all the even elements multiplied by a factor

<!-- grade_id: cell-35937f4621cbea30, points: 0, grade: false, solution: true -->
```python
def multiply_even(input_list, factor):
    return [ x * factor if x % 2 == 0 else x for x in input_list ]
    
multiply_even([1], 10)
```

Output:
```
[1]
```

```python
assert multiply_even([1,2,3,4,5,6,7,8,9],10) == [1, 20, 3, 40, 5, 60, 7, 80, 9]
assert multiply_even([4,2,6,7],10) == [40, 20, 60, 7]
assert multiply_even([],10) == []
assert multiply_even([1],10) == [1]
```

---

## Exercise 3

 use comprehensions to flatten a list of lists 

<!-- grade_id: cell-505e5b66c99b963b, points: 0, grade: false, solution: true -->
```python
l = [[1, 2, 3], [4, 5, 6], [7], [8, 9],[10, 11, 12, 13, 14]]

def flatten_list(l):
    return [ x for subl in l for x in subl ]
```

<!-- grade_id: cell-62297b3504f278ec, points: 1, grade: true, solution: false -->
```python
assert flatten_list([[1, 2, 3], [4, 5, 6], [7], [8, 9],[10, 11, 12, 13, 14]]) == [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14]
assert flatten_list([[1, 2, 3], [4, 5, 6], [7], [8, 9,10, 11, 12, 13, 14]]) == [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14]
assert flatten_list([[1, 2, 3], [4, 5, 6], [7], [8, 9],[10], [11], [12], [13, 14]]) == [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14]
```

---

## Exercise 4

 use comprehensions to filter a list of words keeping only those starting starting with M or containing F

<!-- grade_id: cell-096771dff34a4e5b, points: 0, grade: false, solution: true -->
```python
def filter_list(l: list[str]):
    return [ word for word in l if word.lower().startswith('m') or 'f' in word.lower() ]

```

<!-- grade_id: cell-276004e0334339d7, points: 1, grade: true, solution: false -->
```python
assert filter_list([ 'Monica', 'John', 'Anne', 'Mike', 'Nancy', 'Peter', 'Frank', 'Mary', 'Marc-Frank', 'August-Fracisc'] ) == ['Monica', 'Mike', 'Frank', 'Mary', 'Marc-Frank', 'August-Fracisc']
assert filter_list([ 'a', 'b', 'Fa', 'FM', 'm', 'hmf', 'mMmMfF'] ) == ['Fa', 'FM', 'm', 'hmf', 'mMmMfF']

```

---

## Exercise 5

 use comprehensions to add the first two items of a triple in a list

<!-- grade_id: cell-c3fc3c414ff1fbf1, points: 0, grade: false, solution: true -->
```python
def sum_triples(l):
    return [ a + b for (a, b, _) in l ]

```

<!-- grade_id: cell-5939f57e1805c994, points: 1, grade: true, solution: false -->
```python
assert sum_triples([(1,2,12), (2,1,3), (3,3,1) ] ) == [3, 3, 6]
assert sum_triples([(1,-1,12), (-2,-1,3), (3,3,1) ] ) == [0, -3, 6]
assert sum_triples([(1,21,12), (2,1,3), (3,3,1) ] ) == [22, 3, 6]
```

---

## Exercise 6

 use comprehensions to transpose a squared matrix (i.e., list of lists)

<!-- grade_id: cell-8f9fe87f4d90e397, points: 0, grade: false, solution: true -->
```python
def transpose(mat: list[list[int]]):
    return [ [ mat[j][i] for j in range(len(mat)) ] for i in range(len(mat[0])) ]

print(transpose([
        [0,1,2,3],
        [3,4,5,1],
        [6,7,8,10],
        [9,17,8,110]])
    )
```

Output:
```
[[0, 3, 6, 9], [1, 4, 7, 17], [2, 5, 8, 8], [3, 1, 10, 110]]
```

<!-- grade_id: cell-dd573e16ddec2129, points: 1, grade: true, solution: false -->
```python
assert transpose([
        [0,1,2],
        [3,4,5],
        [6,7,8]]) == [[0, 3, 6], [1, 4, 7], [2, 5, 8]]

assert transpose([
        [0,1,2,3],
        [3,4,5,1],
        [6,7,8,10],
        [9,17,8,110]]) == [[0, 3, 6, 9], [1, 4, 7, 17], [2, 5, 8, 8], [3, 1, 10, 110]]
```

---

## Exercise 7

Define a function `is_prime(n)` that checks whether the number `n` is prime or not (returning `True` or `False` accordingly) and use it to generate the set `p` of all the *reciprocals* (i.e., $1 / n$) of the *prime numbers* smaller than **108**. The set **should** be generated by **set comprehension**.

<!-- grade_id: prime_and_comprehension, points: 0, grade: false, solution: true -->
```python
"""
Function prime
"""
def is_prime(n):
    if n <= 1:
        return False

    for i in range(2, n):
        if n % i == 0:
            return False
        
    return True
```

<!-- grade_id: cell-875e60fc9a8535af, points: 1, grade: true, solution: true -->
```python
"""
Write down the set comprehension expression, storing the data in p
"""
p = { 1 / x for x in range(1, 109) if is_prime(x) }
p
```

Output:
```
{0.009345794392523364,
 0.009708737864077669,
 0.009900990099009901,
 0.010309278350515464,
 0.011235955056179775,
 0.012048192771084338,
 0.012658227848101266,
 0.0136986301369863,
 0.014084507042253521,
 0.014925373134328358,
 0.01639344262295082,
 0.01694915254237288,
 0.018867924528301886,
 0.02127659574468085,
 0.023255813953488372,
 0.024390243902439025,
 0.02702702702702703,
 0.03225806451612903,
 0.034482758620689655,
 0.043478260869565216,
 0.05263157894736842,
 0.058823529411764705,
 0.07692307692307693,
 0.09090909090909091,
 0.14285714285714285,
 0.2,
 0.3333333333333333,
 0.5}
```

<!-- grade_id: correct_prime_and_comprehension, points: 1, grade: true, solution: false -->
```python
"""
Roughly check the prime function
"""
not_primes = [0, 1, 4, 6, 8, 9, 10, 12, 144, 13*13]
for v in not_primes:
    assert not is_prime(v), f"{v} uncorrectly stated as prime"
primes = [2, 3, 5, 7, 7877]
for v in primes:
    assert is_prime(v), f"{v} uncorrectly stated as not prime"
"""
Check that the comprehension set is correct
"""
assert p == {0.009345794392523364,
 0.009708737864077669,
 0.009900990099009901,
 0.010309278350515464,
 0.011235955056179775,
 0.012048192771084338,
 0.012658227848101266,
 0.0136986301369863,
 0.014084507042253521,
 0.014925373134328358,
 0.01639344262295082,
 0.01694915254237288,
 0.018867924528301886,
 0.02127659574468085,
 0.023255813953488372,
 0.024390243902439025,
 0.02702702702702703,
 0.03225806451612903,
 0.034482758620689655,
 0.043478260869565216,
 0.05263157894736842,
 0.058823529411764705,
 0.07692307692307693,
 0.09090909090909091,
 0.14285714285714285,
 0.2,
 0.3333333333333333,
 0.5}
```

```python

```