## General instructions

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel/runtime** (Colab: in the menubar, select *Runtime*$\rightarrow$*Factory Reset Runtime*; Jupyter: in the menubar, select *Kernel*$\rightarrow$*Restart*) and then **run all cells** (Colab: in the menubar, select *Runtime*$\rightarrow$*Run all*; Jupyter: in the menubar, select *Cell*$\rightarrow$*Run All*).

Make sure you fill in any place that says `YOUR CODE HERE` or `"YOUR ANSWER HERE"`, as well as the list of the group members in the following cell.

Enter here the *Group Name* and the list of *Group Members*.

`Doi eletronics`

`Burco Lorenzo, Persello Riccardo`

In order to be able to have an evaluation DO NOT delete/cut the cells with code and answers. Once you have finished you can downolad the notebook (Colab: in the menubar, select *File*$\rightarrow$*Download .ipynb*; Jupyter: in the menubar, select *File*$\rightarrow$*Download as*$\rightarrow$*Notebook (.ipynb)*) and upload as an assignment on the e-learning platform.

The following cell will load the Google Drive extension for the current notebook, when the variable `MOUNT` is `True`. This allow you to mount the Google Drive filesystem for file persistence. The mountpoint will be `/content/gdrive`.
Furthermore, it will set the `PATH` variable, from now on, so that if you have to refer to external files you could do that by writing:

```python
os.path.join(PATH, filename)
```

This will append the filename after the specific PATH.

```python
import os
MOUNT = False
if 'google.colab' in str(get_ipython()) and MOUNT:
    from google.colab import drive
    drive.mount('/content/gdrive')
    PATH = '/content/gdrive/MyDrive'
else:
    PATH = '.'
```

# Important warning

**⚠️ avoid copying, removing or modifying test cells, if you do that your assignment might be graded wrongly ⚠️**

---

Suppose you have the outcomes of an exam reported as string of given answers and correct ones. The format of the data is the following one:

```python
[{'student_id': '1', 'given': ['F'. 'AB', 'C', ...], 'correct': ['F', 'A', 'D', ...] }, ...]
```
That is, there are lists of dictionaries, containing `student_id`, `given` and `correct` answers represented as list of strings. The above example reads `'F'` is the given answer for the first question and `'F'` is the correct answer for the first question. As for the second question `'A'` and `'B'` are the given answer but only `'A'` is correct.

An example file is provided along with this notebook, whose name is `correction.json`.

The following statements will import the library for loading that file and providing its content as the `data` variable. Notice that in order to be able to load it you should **upload it** to the colaboratory environment. Alternatively you might mount the google drive and try referring to its path on the google drive with the `PATH` variable.

```python
import json
import os
if 'PATH' not in locals():
    PATH = '.'

with open(os.path.join(PATH, 'correction.json')) as f:
    data = json.load(f)
```

```python
print(len(data), data[:5])
```

Output:
```
228 [{'student_id': '95', 'correct': ['F', 'A', 'B', 'D', 'C', 'A', 'C', 'C', 'C', 'A', 'D', 'E', 'AB', 'C', 'AD'], 'given': ['F', 'A', 'B', 'A', 'C', 'C', 'C', 'C', 'A', 'C', 'D', 'B', 'AB', 'A', 'CBE']}, {'student_id': '181', 'correct': ['E', 'B', 'A', 'A', 'C', 'B', 'D', 'B', 'B', 'EA', 'B', 'B', 'AE', 'A', 'A'], 'given': ['E', 'B', 'A', 'A', 'C', 'B', 'D', 'B', 'B', 'EBA', 'B', 'B', 'AE', 'C', 'A']}, {'student_id': '60', 'correct': ['C', 'B', 'D', 'C', 'C', 'D', 'E', 'D', 'D', 'B', 'CF', 'D', 'D', 'C', 'A'], 'given': ['C', 'C', 'B', 'D', 'C', 'E', 'E', 'D', 'D', 'D', 'D', 'BD', 'B', 'B', 'BC']}, {'student_id': '50', 'correct': ['B', 'A', 'B', 'B', 'AD', 'A', 'B', 'B', 'B', 'B', 'A', 'C', 'D', 'A', 'B'], 'given': ['B', 'A', 'B', 'B', 'A', 'A', 'B', 'B', 'B', 'B', 'A', 'C', 'D', 'A', 'B']}, {'student_id': '55', 'correct': ['B', 'B', 'C', 'C', 'D', 'D', 'ADFB', 'B', 'D', 'A', 'B', 'C', 'C', 'C', 'B'], 'given': ['B', 'B', 'C', 'D', 'D', 'D', 'ADFB', 'B', 'D', 'A', 'B', 'C', 'C', 'C', '
```

## Exercise 1

Write a function named `correction` that takes as arguments the answers to a single question and returns the following values:

1. the number of correct answers among the given ones;
2. the number of missing answers from the given ones.

For example, in the case of `'AB'` as the given answers and `'BD'` as the correct ones the only correct answer in the given is `B` whereas `D` is missing.

**Suggestion:** use the “super-powers” of the **`set`** data structure.

<!-- grade_id: cell-c2760e69b1cd2873, points: 0, grade: false, solution: true -->
```python
def correction(given: str, correct: str) -> (int, int):
    given_set = set(given)
    correct_set = set(correct)

    intersection = given_set & correct_set
    missing = correct_set - intersection

    return (len(intersection), len(missing))
```

<!-- grade_id: cell-7971e816718e146d, points: 1, grade: true, solution: false -->
```python
"""
Checks that the function `correction` exists
"""
try:
    correction
except NameError:
    assert False, "The name correction does not exists"
"""
A few cases
"""
assert correction('AB', 'BD') == (1, 1)
assert correction('A', 'BDC') == (0, 3)
assert correction('', '') == (0, 0)
assert correction('', 'A') == (0, 1)
```

## Exercise 2

Write a function `process` that transforms the data from the list of student into the following new format:

```python
[{ 'student_id': '1', 'answers': [(1, 1), (0, 3), ... ] }, { 'student_id': 2, ... }]
```

by using the `correction` function defined above. The first element of the tuple should be the number of correct answers and the second element of the tuple the number of missing ones.

Optionally, the function `process` could take further arguments `given` and `missing` that, when set to `True`, will fill in the `answers` list with the specified value. That is:

```python
process(data) # should perform the above described transformation, with a tuple of two values
process(data, given=True, missing=True) # same as before
process(data, given=True) # should fill in `answers` list with just the `given` values
process(data, missing=True) # should fill in `answers` list with just the `missing` values
```

<!-- grade_id: cell-78cf6f8aac913681, points: 0, grade: false, solution: true -->
```python
def split_correction(correction_result: (int, int), given, missing):

    if not given and not missing:
        # No parameters have been passed, keep the result.
        return correction_result
    if given and missing:
        # Same as before.
        return correction_result
    elif missing:
        return correction_result[1]
    elif given:
        return correction_result[0]

def process(data, given=False, missing=False):

    result = []
    for student_data in data:
        result.append(
            { 
              'student_id': student_data['student_id'],
              'answers': [ split_correction(correction(given_answer, correct_answer), given, missing) for (given_answer, correct_answer) in zip(student_data['given'], student_data['correct'])]
            }
        )

    return result
```

<!-- grade_id: cell-880f468f6c45c067, points: 0, grade: true, solution: false -->
```python
res = process(data)
for i, d in enumerate(res):
    assert d['student_id'] == data[i]['student_id'], f"{d['student_id']} {data[i]}"
    for t, g, c in zip(d['answers'], data[i]['given'], data[i]['correct']):
        assert t == correction(g, c)
res = process(data, given=True)
for i, d in enumerate(res):
    assert d['student_id'] == data[i]['student_id'], f"{d['student_id']} {data[i]}"
    for t, g, c in zip(d['answers'], data[i]['given'], data[i]['correct']):
        assert t == correction(g, c)[0]
res = process(data, missing=True)
for i, d in enumerate(res):
    assert d['student_id'] == data[i]['student_id'], f"{d['student_id']} {data[i]}"
    for t, g, c in zip(d['answers'], data[i]['given'], data[i]['correct']):
        assert t == correction(g, c)[1]
```

## Exercise 3

Write a function `report(data, f)` that takes as argument the list obtained by the `process` function and an aggregation function `f` (can be `min`, `max`, `sum`).

The `report` function should take the `answers` list of each student and aggregate the **number of correct** results by applying the aggregation function `f`.

The default value for the aggregation function for the `report` function should be `sum`.

The output should be a dictionary, whose keys are the student ids and the values are the result of the aggregation, for example:

```python
{'95': 8,
 '181': 1,
 '60': 10,
 '50': 1,
 ...
}
```

<!-- grade_id: cell-5eb54861aa4f2e39, points: 0, grade: false, solution: true -->
```python
def report(data, f=sum):
    return { student_data['student_id']: f([ num_correct for (num_correct, _) in student_data['answers']]) for student_data in data }
```

<!-- grade_id: cell-f845536578dd20bf, points: 1, grade: true, solution: false -->
```python
# This will check the results with respect to the reference ones, provided in a json file
# In order to pass the test also this file should be uploaded
res = process(data)
with open(os.path.join(PATH, 'results.json')) as f:
    compare = json.load(f)
for student_id, value in report(res).items():
    assert value == compare[student_id]
```

## Exercise 4

Write an `average(d)` function that given a list of numerical values returns the average value of the list and use it as the *aggregator* in the `report()` function defined above.

<!-- grade_id: cell-8a1f3d9b03950928, points: 0, grade: false, solution: true -->
```python
def average(d):
    s = sum(d)
    return s / len(d)
```

<!-- grade_id: cell-6b48d9a6ce4908ec, points: 1, grade: true, solution: false -->
```python
try:
    report(res, average)
except NameError as e:
    assert False, str(e)
```

```python

```