## General instructions

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel/runtime** (Colab: in the menubar, select *Runtime*$\rightarrow$*Factory Reset Runtime*; Jupyter: in the menubar, select *Kernel*$\rightarrow$*Restart*) and then **run all cells** (Colab: in the menubar, select *Runtime*$\rightarrow$*Run all*; Jupyter: in the menubar, select *Cell*$\rightarrow$*Run All*).

Make sure you fill in any place that says `YOUR CODE HERE` or `"YOUR ANSWER HERE"`, as well as the list of the group members in the following cell.

Enter here the *Group Name* and the list of *Group Members*.

`Doi eletronics`

`Burco Lorenzo, Persello Riccardo`

In order to be able to have an evaluation DO NOT delete/cut the cells with code and answers. Once you have finished you can downolad the notebook (Colab: in the menubar, select *File*$\rightarrow$*Download .ipynb*; Jupyter: in the menubar, select *File*$\rightarrow$*Download as*$\rightarrow$*Notebook (.ipynb)*) and upload as an assignment on the e-learning platform.

The following cell will load the Google Drive extension for the current notebook, when the variable `MOUNT` is `True`. This allow you to mount the Google Drive filesystem for file persistence. The mountpoint will be `/content/gdrive`.
Furthermore, it will set the `PATH` variable, from now on, so that if you have to refer to external files you could do that by writing:

```python
os.path.join(PATH, filename)
```

This will append the filename after the specific PATH.

```python
import os
MOUNT = False
if 'google.colab' in str(get_ipython()) and MOUNT:
    from google.colab import drive
    drive.mount('/content/gdrive')
    PATH = '/content/gdrive/MyDrive'
else:
    PATH = '.'
```

# Important warning

**⚠️ avoid copying, removing or modifying test cells, if you do that your assignment might be graded wrongly ⚠️**

---

# The very basics

In this notebook we are dealing with the real basics of the Python language, trying to cope with variables and performing some stuff.


## Exercise 1

Write the code that, given a value stored in a variable `n` it checks whether that number is a prime or not and stores the truth value of that condition in another variable `prime`.

**Recall**: a natural number (i.e., integer and non negative) is *prime* if it is only divisible by 1 and itself.

<!-- grade_id: cell-3dd7f008193cd592, points: 0, grade: false, solution: true -->
```python
n = 6
for i in range(2, n):
    if n % i != 0:
        prime = True
    else:
        prime = False
        break

print(prime)
```

Output:
```
False
```

<!-- grade_id: cell-40afce5bc2d3ca0c, points: 1, grade: true, solution: false -->
```python
assert prime is False
```

Try another case of running the above code (you have to copy & paste it in the following cell) for testing another number that, instead, is prime (that is, copy the code and change the value of `n` to test a prime number).

<!-- grade_id: cell-299112a0e5a1f839, points: 0, grade: true, solution: true -->
```python
n = 7
for i in range(2, n):
    if n % i != 0:
        prime = True
    else:
        prime = False
        break

print(prime)
```

Output:
```
True
```

<!-- grade_id: cell-ff6c2d61bcd19008, points: 0, grade: true, solution: false -->
```python
assert prime is True
```

For the purpose of reuse, the code you wrote can be embedded in a *function*. More details about functions will follow in a later video lecture, although by now it's enough to know that a function can be defined as follows:

```python
def function_name(parameter1, parameter2): # the function HAS parameters, which represent data coming from the outside environment
    # function body, i.e., the code that implements the function and compute a `value`
    return value # this is usually the final statement, passing the result of the function to the outside environment
```

Once defined, the function can be called multiple times as:

```python
result = function_name(5, 2)
```

and the returned value can be stored, e.g., in variables.

## Exercise 2

Embed the code you write above in a function, named `prime()` that takes the value for which to determine its primality as a parameter.

Notice that the name `prime` now has become the name of the function and cannot be used as a variable inside the function body (this means that you cannot just copy and paste the above code but rather you must adapt it to the specific situation).

<!-- grade_id: cell-35c94d1e8cdec97f, points: 0, grade: false, solution: true -->
```python
def prime(n):
    for i in range(2, n):
        if n % i == 0:
            return False
        
    return True
```

Notice that the name `prime` now has changed its nature: from the previous cell on it has become a `function`. You can inspect it by adding a new cell after this one and issuing the command `type(prime)` which tells you which is the type associated to the variable.

You can also try re-running the previous cells in order to check what was the type of `prime` earlier.

<!-- grade_id: cell-3c90c7b21e1eaba1, points: 1, grade: true, solution: false -->
```python
assert prime(6) is False
assert prime(5) is True
assert prime(45) is False
```

## Excercise 3

The value of the stopped factorial `stopped_fact(n, k)` function consists in the multiplication of the numbers $n$, $n - 1$, $n - 2$, ..., $n - k + 1$. For example `stopped_fact(5, 3)` has value $5 \times 4 \times 3 = 60$, or `stopped_fact(13, 4)` has value $13 \times 12 \times 11 \times 10 = 17160$.

Fill in the function to compute the stopped factorial. 

<!-- grade_id: cell-9c9ca1ea16bf3107, points: 0, grade: false, solution: true -->
```python
def stopped_fact(n, k):
    acc = 1
    for x in range(n-(k-1), n+1):
        acc *= x

    return acc
```

<!-- grade_id: cell-088cedc90fdc6551, points: 1, grade: true, solution: false -->
```python
assert stopped_fact(4, 2) == 12
assert stopped_fact(5, 3) == 60
assert stopped_fact(13, 4) == 17160
```

```python

```